```python
# Initialize Otter
import otter
grader = otter.Notebook("DSA_2023_Entry.ipynb")
```

## DSA 2023 Summer School Admittance Check

Thanks for your interest in attending DSA Kigali 2023. To attend the summer school you have to have some level of basic Python proficiency. Completing the following notebook should ensure you have the right kind of background to benefit maximally from the Summer School. See you in Kigali!

```python
# Run these once ... Just in case 
#!pip install nose 
#!pip install otter-grader
import IPython
from IPython import get_ipython
# Import the good stuff
import pandas as pd
import numpy as np
import math
from nose.tools import assert_equal
import otter
grader = otter.Notebook()
```

**Question 1:** Write a Python function to return a tuple of primes and non primes given an integer input 

```python
...
num = 30
print("Following are the prime and non prime numbers smaller than or equal to", num)
print(sieve(num))
```

```python
grader.check("q1")
```

**Question 2:** Create a dictionary to store the type and number of unique characters in the paragraph provided

```python
paragraph = "The following is sample text I used to practice special characters using keybr.com: 112233445566778899 Saturn V rocket’s first stage carries 203,400 gallons (770,000 liters) of kerosene fuel and 318,000 gallons (1.2 million liters) of liquid oxygen needed for combustion. At liftoff, the stage’s five F-1 rocket engines ignite and produce 7.5 million pounds of thrust. To replace those goofy quantities with the far less retarded metric system (even though liters are considered part of the metric system they are the same as cubic deci-meters) you would say 770 cubic meters of kerosene {abbreviated as m3} and 1,204 m3 of liquid O2 [O2 is the symbol for oxygen]. We would also say it produced 33,600,000 newtons of force <abbreviated as N>. To add scientific notation {a way of writing numbers that allows you to write only as many digits `of specificity` as you would like} you can write 7.7 * 10 ^ 2 m3 of kerosene 1.204 * 10 ^ 3 m3 of O2 and 3.3 * 10 ^ 7 newtons. Another way to write scientific notation is to replace the “* 10 ^” with ‘E’ -/capital e\-. So our numbers would look like:s"
...
print(dictionary)
```

```python
grader.check("q2")
```

**Question 3:** Extract the values from the dictionary above into a list

```python
...
print (keysList)
```

```python
grader.check("q3")
```

**Question 4:** Extract the keys from the dictionary above into a list

```python
...
print (valuesList)
```

```python
grader.check("q4")
```

**Question 5:** Merge the two lists into a tuple

```python
...
print(final_tuple)
```

```python
grader.check("q5")
```

**Question 6:**  Write a function `greatest_common_divisor` that takes two inputs `a` and `b` and returns the greatest common divisor of the two numbers. E.g. input `(10, 15)` would return `5`

```python
def greatest_common_divisor(a, b):
    ...
```

```python
grader.check("q6")
```

**Question 7:**  Write a function `get_nearest_farthest` that takes in a point of interest (pt) and a list of points and returns the nearest point and the farthest point from the point of intest.

```python
def get_nearest_farthest(pt, ptlist):
    ...
```

```python
grader.check("q7")
```

**Question 8:**  Write a function `perfectly_divisible` to return a list of numbers between 0 and a number $N$ that are perfectly divisible by $q$ (with out leaving a remainder). <br>
**Hint**: $N$ should also be inclusive in the numbers being considered. If N is negative use N == 20

```python
def perfectly_divisible(N, q):
    ...
```

```python
grader.check("q8")
```

**Question 9:**  Write a function `flatten_lists` that takes in a list of lists and outputs a sorted list of elements of sublists of the input list (confusing right?) <br>
Example: given `flatten_lists([[2,13,44], [6,7]])` it should return `[2,6,7,13,44]`

```python
def flatten_lists(superlist):
    ...
```

```python
grader.check("q9")
```

**The Extra Mile!** 

**Download the dataset "Nakuru Sensor Data for February 2023 from https://open.africa/dataset/sensorsafrica-airquality-archive-nakuru/resource/14165682-37dd-4f40-914e-eb24619e4ef8"**

Load the 'timestamp','value_type' and 'value' collumns into a pandas dataframe 

```python
import pandas as pd
from functools import reduce

#Load the 'timestamp','value_type' and 'value' collumns into a pandas dataframe 

df = pd.read_csv(r'Nakuru_Sensor_Data_Feb_2023.csv',usecols=['timestamp','value_type','value'])

print(df)
```

**Question E1:** Create a new collumn 'timestamp_new' from 'timestamp' rounded off to the second 

```python
...
print(df)
```

**Question E2:** Split the dataframe based on the 'value_type' collumn

```python
grouped = ...
df_temperature = ...
df_humidity = ...
df_P1 = ...
df_P2 = ...
```

**Question E3:** Find the mean, mode and median of the humidity, temperature, P1 and P2 values

```python
...
print("Temperature")
print("Mean = ",mu_temp, "Median = ", med_temp, "Standard Deviation = ", sd_temp)

...
print("Humidity")
print("Mean = ",mu_humidity, "Median = ", med_humidity, "Standard Deviation = ", sd_humidity)

...
print("P1")
print("Mean = ",mu_P1, "Median = ", med_P1, "Standard Deviation = ", sd_P1)

...
print("P2")
print("Mean = ",mu_P2, "Median = ", med_P2, "Standard Deviation = ", sd_P2)

```

**Question E4:** Merge the grouped dataframes into one dataframe and save it as newdata.csv

```python
...
print(df_merged)

from pathlib import Path  
filepath = Path('newdata.csv')  
filepath.parent.mkdir(parents=True, exist_ok=True)  
df_merged.to_csv(filepath)  
```

**Question E5:** Open the new data file and select only 'timestamp_new','value_temperature','value_humidity','value_P1' and 'value_P2' collunms

```python
...
new_df 
```

**Question E6:** Replace the missing values in collumns 'value_P1' and 'value_P2' with the mean of the same collumns

```python
#Check for missing values 
missing_values = ...

# displaying data only with NaN
...
```

```python
#print("The mean of P1",new_df['value_P2'].mean())
#print("The mean of P2",new_df['value_P1'].mean())

#replace missing values with the mean 
...
new_df
```

**Question E7:** Compute the correlations between temperature and humidity; P1 and P2 

```python
#Correlation between temperature and humidity
...

print ("Correlation between temparature and humidity",corr_temp_hum)

#Correlation between P1 and P2
...

print ("Correlation between P1 and P2",corr_p1_p2)
```

**Question E8:** Create a new dataframe comprising of the average values of the variables by day

```python
...
means_feb = ...
means_feb
```

**Question E9:** Download the dataset "Nakuru Sensor Data for November 2021 from https://open.africa/dataset/sensorsafrica-airquality-archive-nakuru/resource/1f11351e-156b-4e4c-a063-b86a6ee07c4a"

Load the data into a pandas dataframe and develop a dataframe similar to what you have just done in Question no's 6 - 13 above

```python
...
means_nov = ...
means_nov
```

**Question E10:** Convert the means_feb and means_nov to numpy arrays ... and then compute the average of each collumn

Hint: The results are numpy arrays

```python
feb_numpy_array = ...
feb_average = ...
feb_average
```

```python
nov_numpy_array = ...
nov_average = ...
nov_average
```

**Question E11:** Write a function to compute the difference in the average temperature, humidity, P1 and P2 between November 2021 and February 2023

Hint: The result is a vector of the form [w,x,y,z]

```python
def sub():
    ...
    return diff
 
# Print the result of the subtraction
difference = sub()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Upload the generated zip file to the Google sheet please. Uploading work which is not yours will lead to non admittance.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: 2)
```
>>> # BEGIN TEST CONFIG
>>> points: 2
>>> # END TEST CONFIG
>>> def test_primes(sieve):
...     assert sieve(30) == ((2, 3, 5, 7, 11, 13, 17, 19, 23, 29), (4, 6, 8, 9, 10, 12, 14, 15, 16, 18, 20, 21, 22, 24, 25, 26, 27, 28, 30))
>>> test_primes(sieve)
```


Tests for q2 (points: None)
```
>>> # BEGIN TEST CONFIG
>>> points: 2
>>> # END TEST CONFIG
>>> def test_dictionary(Counter):
...     assert Counter(paragraph) == {' ': 192, 'e': 89, 'o': 68, 't': 65, 'i': 60, 's': 57, 'a': 48, 'r': 45, 'n': 42, 'l': 35, 'c': 31, 'd': 27, 'f': 26, 'u': 26, '0': 22, 'm': 20, 'y': 19, 'h': 17, 'w': 16, 'g': 14, '3': 13, '.': 12, 'p': 11, 'b': 11, '1': 11, '2': 10, '7': 10, 'k': 9, ',': 7, '4': 5, 'v': 4, '-': 4, '*': 4, '^': 4, 'T': 3, 'x': 3, '5': 3, '6': 3, '8': 3, '’': 3, '(': 3, ')': 3, 'q': 3, 'O': 3, ':': 2, '9': 2, 'S': 2, 'A': 2, '{': 2, '}': 2, '`': 2, 'I': 1, 'V': 1, 'F': 1, '[': 1, ']': 1, 'W': 1, '<': 1, 'N': 1, '>': 1, '“': 1, '”': 1, '‘': 1, 'E': 1, '/': 1, '\\': 1}
```


Tests for q3 (points: None)
```
>>> # BEGIN TEST CONFIG
>>> points: 2
>>> # END TEST CONFIG
>>> def test_keyslist():
...     assert list(dictionary.keys()) == ['T', 'h', 'e', ' ', 'f', 'o', 'l', 'w', 'i', 'n', 'g', 's', 'a', 'm', 'p', 't', 'x', 'I', 'u', 'd', 'r', 'c', 'k', 'y', 'b', '.', ':', '1', '2', '3', '4', '5', '6', '7', '8', '9', 'S', 'V', '’', '0', ',', '(', ')', 'q', 'A', 'v', 'F', '-', '{', '}', 'O', '[', ']', 'W', '<', 'N', '>', '`', '*', '^', '“', '”', '‘', 'E', '/', '\\']
>>> test_keyslist()
```


Tests for q4 (points: None)
```
>>> # BEGIN TEST CONFIG
>>> points: 2
>>> # END TEST CONFIG
>>> def test_valueslist():
...     assert list(dictionary.values()) == [3, 17, 89, 192, 26, 68, 35, 16, 60, 42, 14, 57, 48, 20, 11, 65, 3, 1, 26, 27, 45, 31, 9, 19, 11, 12, 2, 11, 10, 13, 5, 3, 3, 10, 3, 2, 2, 1, 3, 22, 7, 3, 3, 3, 2, 4, 1, 4, 2, 2, 3, 1, 1, 1, 1, 1, 1, 2, 4, 4, 1, 1, 1, 1, 1, 1]
```


Tests for q5 (points: None)
```
>>> # BEGIN TEST CONFIG
>>> points: 2
>>> # END TEST CONFIG
>>> def test_tuple():
...     assert final_tuple == (('T', 'h', 'e', ' ', 'f', 'o', 'l', 'w', 'i', 'n', 'g', 's', 'a', 'm', 'p', 't', 'x', 'I', 'u', 'd', 'r', 'c', 'k', 'y', 'b', '.', ':', '1', '2', '3', '4', '5', '6', '7', '8', '9', 'S', 'V', '’', '0', ',', '(', ')', 'q', 'A', 'v', 'F', '-', '{', '}', 'O', '[', ']', 'W', '<', 'N', '>', '`', '*', '^', '“', '”', '‘', 'E', '/', '\\'), (3, 17, 89, 192, 26, 68, 35, 16, 60, 42, 14, 57, 48, 20, 11, 65, 3, 1, 26, 27, 45, 31, 9, 19, 11, 12, 2, 11, 10, 13, 5, 3, 3, 10, 3, 2, 2, 1, 3, 22, 7, 3, 3, 3, 2, 4, 1, 4, 2, 2, 3, 1, 1, 1, 1, 1, 1, 2, 4, 4, 1, 1, 1, 1, 1, 1))
>>> test_tuple()
```


Tests for q6 (points: 10)
```
>>> assert_equal(greatest_common_divisor(10, 15), 5)
```


Tests for q7 (points: 10)
```
>>> assert_equal(get_nearest_farthest((3, 8), [(9, 3), (8,5), (7,6)]), ((7, 6), (9, 3)))
```


Tests for q8 (points: 5)
```
>>> assert_equal(perfectly_divisible(10,2), [0,2,4,6,8,10])
```


Tests for q9 (points: 4)
```
>>> assert_equal(flatten_lists([[2,13,44], [6,7]]), [2, 6, 7, 13, 44])
```
